#### Question 1
Ecrivez une fonction python qui prend en argument une chaine de caractères et remplace
toutes les occurrences de la 1ere lettre par le symbole & excepté la 1ere lettre.
Exemple de chaine de caractères : 'ananas'
Résultat attendu : 'an&n&s'

<!-- grade_id: cell-c10287dc96becd68, points: 0, grade: false, solution: true -->
```python
def repl(char):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-fb28f65e1f4d7017, points: 2, grade: true, solution: false -->
```python
assert repl('ananas')=='an&n&s'
assert repl('jojo')=='jo&o'
assert repl('abel')=='abel'
```

#### Question 2
Ecrivez une fonction Python qui prend une liste de mots et renvoie la longueur du plus long

<!-- grade_id: cell-fc464bab5c376def, points: 0, grade: false, solution: true -->
```python
def get_len(ls):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d28e1d57b6ca0b27, points: 2, grade: true, solution: false -->
```python
assert get_len(['abel','gbeli','jonathan','josias'])==8
```

#### Question 3
Ecrivez une fonction Python qui pour une chaine de caractères donnée permute le 1er et le
dernier caractère

<!-- grade_id: cell-316324b6d1a9b33d, points: 0, grade: false, solution: true -->
```python
def permut(char):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b2165cd064b51dea, points: 2, grade: true, solution: false -->
```python
assert permut('abel')=='lbea'
assert permut('label')=='label'
```

#### Question 4
Ecrivez un programme python qui fait la somme de tous les éléments d’une liste

<!-- grade_id: cell-e8ecdbefa90d761b, points: 0, grade: false, solution: true -->
```python
def add(ls):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a80c01dd53569c2d, points: 2, grade: true, solution: false -->
```python
assert add([3,2,9])==14
assert add([0,-2,2])==0
```

#### Question 5
Ecrivez un programme python qui donne le plus grand nombre et le plus petit nombre d’une
liste

<!-- grade_id: cell-d9390c4867c39f7a, points: 0, grade: false, solution: true -->
```python
def min_max(ls):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6db871ee2ee26098, points: 2, grade: true, solution: false -->
```python
assert min_max([3,2,0])==(3,0)
assert min_max([-1,-1])==(-1,-1)
```